# Programming Assignment 4

## Objectives 

The purpose of this assignment is to: 

- practice creating functions by creating two functions and integrating them into a program

## Assignment Instructions

Write two functions, described below.  Be sure to test your code against the test cases!  Submit your code once your tests are passing.  (It is possible to earn partial credit for passing some of the tests, but not passing others.) 

## Let's start coding!

### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

PLEASE BE SURE TO EXECUTE THE FOLLOWING CELL, TO INSTALL THE LIBRARIES NEEDED FOR THIS LAB.  
Note: You will encounter errors if you skip this step.

```python
# Install Necessary Dependencies
import nose, inspect
from nose.tools import assert_equal
```

### Function 1: Prompt the user to select from a list of menu options

Define a function that will take no arguments, but will return an integer that represents the option selected (option 1, 2, 3, or 4) or return a 0 if there was an error.  

The "inside" of this function should do several things, and use one input() statement, and a longer if-elif-else statement.

First, print a list of menu options that the user can select from. Use these 4 options, with the **exact** text below:

1. See an uplifting message
2. See a programming tip
3. Watch me count
4. Quit this program.  

Then, ask the user to enter a number to indicate their selection.

Then, use an if statement to see if they typed a 1.  If so, **return the integer 1**.  

Similarly, use elif statements to **return 2, 3, and 4 as integers** for those options selected.

Finally, include an else statement that tells them that they entered an invalid number.  **Return a 0 in that case.**

#### *Hint:  you can compare strings, like this:*

```python
if user_input == '1': 
    # they entered a 1, great!
```

In this way, you avoid the program crashing if they type "asdf" and then you try to convert that to an integer.

####

<!-- grade_id: cell-fb4aeb0d1c5c8ec4, points: 0, grade: false, solution: true -->
```python
def menu():
    ''' print a list of menu options, and return an integer representing the user's choice '''
    print("1. See an uplifting message")
    print("2. See a programming tip")
    print("3. Watch me count")
    print("4. Quit this program.")
    
    user_input = input("Enter a number to indicate your selection: ")
    
    if user_input == '1':
        return 1
    elif user_input == '2':
        return 2
    elif user_input == '3':
        return 3
    elif user_input == '4':
        return 4
    else:
        print("Invalid number. Please enter a number between 1 and 4.")
        return 0

    raise NotImplementedError # delete this line after you have written your function
```

The cell below allows you to test your function – run it multiple times, trying all sorts of different options, and see that the right integer is returned for each case.

```python
# Manual testing of menu()
this_result = menu()
print('The menu function returned a', this_result, '. Does that seem correct?')
```

Output:
```
1. See an uplifting message
2. See a programming tip
3. Watch me count
4. Quit this program.
Enter a number to indicate your selection:  1
The menu function returned a 1 . Does that seem correct?
```

Now, run the unit tests that the autograder uses, and see how your code works!

<!-- grade_id: cell-b72fe6462a70755e, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASE 1: This cell should not deliver any errors when this cell is run
# Check - for loop implementation
if_statement_found = False
option_input_found = False
see_uplifiting_message_found = False
see_programming_tip_found = False
watch_me_count_found = False
quit_this_program_found = False

menu_lines = inspect.getsource(menu)
code_lines = menu_lines.splitlines()
for line in code_lines:
    if "if" in line:
        if_statement_found = True
    if "input" in line:
        option_input_found = True
    if "see an uplifting message" in line.lower():
        see_uplifiting_message_found = True
    if "see a programming tip" in line.lower():
        see_programming_tip_found = True
    if "watch me count" in line.lower():
        watch_me_count_found = True
    if "quit this program" in line.lower():
        quit_this_program_found = True
        
assert if_statement_found == True, "Please verify you've implemented if logic in your function."
assert option_input_found == True, "Please verify if you've asked the user for their input."
assert see_uplifiting_message_found == True, "Please verify that you have implemented all the menu options, using the EXACT test requested, for example ‘see an uplifting message’ should appear in your menu output"
assert see_programming_tip_found == True, "Please verify that you have implemented all the menu options, using the EXACT test requested, for example ‘see a programming tip’ should appear in your menu output"
assert watch_me_count_found == True, "Please verify that you have implemented all the menu options, using the EXACT test requested, for example ‘watch me count’ should appear in your menu output"
assert quit_this_program_found == True, "Please verify that you have implemented all the menu options, using the EXACT test requested, for example ‘quit this program’ should appear in your menu output"
```

### Function 2: Count up to a number

Write a function with one argument that is an integer.  That function should use a for loop to print out the numbers from 1 to that number (inclusive).

<!-- grade_id: cell-95f796f66fb927cc, points: 0, grade: false, solution: true -->
```python
def count_up(num):
    ''' count from 1 to that number, inclusive '''
    for i in range(1, num + 1):
        print(i)
```

The cell below allows you to test your function – run it multiple times, chaning the value stored in the variable test_value, and see if the function counts correctly for each case

```python
# Manual testing of count_up()
test_value = 5
count_up(test_value)
print('Did your function print the numbers starting a 1 and ending by printing a', test_value, '?')
```

Output:
```
1
2
3
4
5
Did your function print the numbers starting a 1 and ending by printing a 5 ?
```

Now, run the unit tests that the autograder uses, and see how your code works!

<!-- grade_id: cell-a296d0b0289e0378, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASE 2: This cell should not deliver any errors when this cell is run

# Check - for loop implementation
for_loop_found = False

count_up_lines = inspect.getsource(count_up)
code_lines = count_up_lines.splitlines()
for line in code_lines:
    if "for" in line and "in" in line and ":" in line:
        for_loop_found = True

assert for_loop_found == True, "Please verify you've implemented a for loop in your count_up function"
assert count_up(5) == None, "Check your count_up function with number 5."
```

Output:
```
1
2
3
4
5
```

### Main Program: *Putting it all together*

Now pull these two functions together into a whole program.  Your code should use your first function to present the user with a menu of options.  
If your first function returns a 1, your program will print an uplifting message (choose a quote that you find encouraging).  
If they instead pressed a 2, your program will print a programming tip (choose something that you have learned in this class).  
If they pressed 3, your program will then ask them for a number, then count up to that number, starting at 1 (using the function that you wrote).  
If they press 4, the program will exit.   

In this case, you do not need to continue looping – the user would need to run the program another time in order to choose a different option. 

When you run your program, it should produce output similar to what is shown below.  It does not need to be identical, but should be functionally the same. 

```python
def main():
    # to start off, prompt the user to choose an option by using your menu() function
    request = menu()

    # now, use an if/elif/else statement to handle each of the possible options
    # ADD YOUR CODE HERE
    

# This next line lets you run your code by calling the main() function.  Test your work! Note that this section is not being graded, because using the autograder here would force us to add many constraints and reduce the freedom that you have to solve the problem
main()
```

Output:
```
1. See an uplifting message
2. See a programming tip
3. Watch me count
4. Quit this program.
Enter a number to indicate your selection:  4
```