# Programming Assignment 3b

## Objectives

The purpose of this assignment is to:
* practice defining (creating) a function, in particular to create a function that takes one argument and returns a value 
* practice calling (using) a function 
* develop a more complicated program that also uses a while loop to interact with the user repeatedly 

## Assignment Instructions

Write two functions, described below. Be sure to test your code against the test cases! Submit your code once your tests are passing.  

*Note: It is possible to earn partial credit for passing some of the tests, but not passing others.*

## Let's start coding!

#### Using this Notebook

Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE. If you have written the functions correctly, you should see no errors.

When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

PLEASE BE SURE TO EXECUTE THE FOLLOWING CELL,TO INSTALL THE LIBRARIES NEEDED FOR THIS LAB.  
Note: You will encounter errors if you skip this step.

```python
# Install Necessary Dependencies
import math, nose
from nose.tools import assert_almost_equal
```

### __Function 1__:  Area of a Circle

Write a function to calculate the area of a circle, using the formula `pi*r*r`.  If you want, you can use the constant math.pi instead of entering 3.14159.  

<!-- grade_id: cell-f65fe616cbb0881f, points: 0, grade: false, solution: true -->
```python
### GRADED FUNCTION 1: Write a function to calculate the area of a circle.
def area(radius):
    ''' return the area of a circle of this radius '''
    return math.pi * radius**2
    raise NotImplementedError()
```

Now test your code to see that it works the way that you want it to! The cell below gives you an opportunity to test your code in a way similar to how the auto-grader will test it.  The first line of code allows you to replace the 10 with a number of your choice.  When you run the cell, the second line will call the function that contains your solution.  The third line then prints the results, so that you can manually test it – if you put these numbers into a calculator, are they correct?  Try it for different values.

```python
# Manually test `area`
test_radius_for_area = 10 # replace 10 with an appropriate value of your choice for the radius variable
test_area_function = area(test_radius_for_area)
print(f"radius of circle: {test_radius_for_area}, calculated area: {test_area_function}")
```

Output:
```
radius of circle: 10, calculated area: 314.1592653589793
```

The next cell shows how the auto-grading tests are executed.  Several possible cases are run through the function, and the results are tested against that actual correct answer. If you run the test cell and there is no output, then the tests all passed – Great!

If you’re curious, the function named assert_almost_equal() is a way to test floating point numbers – if you remember from the lectures, these are always rounded at the end, and so we cannot say that they are exactly equal, but only that they are pretty close to the same number. 

<!-- grade_id: cell-3ada116f80055b02, points: 5, grade: true, solution: false -->
```python
# TEST CASES: This cell should not deliver any errors if all tests are passing.
assert_almost_equal(area(1), math.pi, places=2, msg="Check your area function.")
assert_almost_equal(area(0), 0, places=2, msg="Check your area function with radius 0.")
assert_almost_equal(area(2), 4*math.pi, places=2, msg="Check your area function with radius 2.")
assert_almost_equal(area(10), 100*math.pi, places=2, msg="Check your area function with radius 10.")
```

### Function 2: Circumference

Write a function to calculate the circumference of a circle, using the formula `2*pi*r`. If you want, you can use the constant math.pi again.

<!-- grade_id: cell-97bec1841199b27e, points: 0, grade: false, solution: true -->
```python
def circumference(radius):
    ''' return the circumference of a circle of this radius '''
    return 2 * math.pi * radius
    raise NotImplementedError()
```

Again, run your function a few times to test it with different values.  Does your result match what you would find using a calculator?

```python
# Manually test `circumference`
test_radius_for_circumference = 10 # replace 10 with an appropriate value of your choice for the radius variable
test_area_function = circumference(test_radius_for_circumference)
print(f"radius of circle: {test_radius_for_circumference}, calculated area: {test_area_function}")
```

Output:
```
radius of circle: 10, calculated area: 62.83185307179586
```

Now, run the unit test to see how the autograder will test your function.

<!-- grade_id: cell-0cead43d8902123f, points: 5, grade: true, solution: false -->
```python
# TEST CASES: This cell should not deliver any errors if all tests are passing.
assert_almost_equal(circumference(1), 2*math.pi, places=2, msg="Check your circumference function.")
assert_almost_equal(circumference(0), 0, places=2, msg="Check your circumference function with radius 0.")
assert_almost_equal(circumference(2), 4*math.pi, places=2, msg="Check your circumference function with radius 2.")
assert_almost_equal(circumference(10), 20*math.pi, places=2, msg="Check your circumference function with radius 10.")
```

## Putting it all together

Now let's put it all together! Below is a program that will prompt the user (you in this case!) for two measurement values. 
The program will then use these measurements to calculate an area and circumference of a circle. You can try this out below and watch how your functions come to life in a real-world application!  Feel free to add or modify this code – this cell is not graded, but is a great way to learn more about how functions work.

```python
def main():
    
    while True:
        # ask if they want an area or circumference
        answer = input('Enter 1 for the area\nEnter 2 for the circumference\nEnter 0 to exit\n')

        # do what the user says
        if answer == '1':  # area
            # prompt user for radius
            rad = float(input('Enter a radius: '))
            print(f'The area is {area(rad):.2f}')
        elif answer == '2': # circumference
            # prompt user for radius
            rad = float(input('Enter a radius: '))
            print(f'The circumference is {circumference(rad):.2f}')
        elif answer == '0': # quit
            break
        else: # they entered something else|
            print('I do not recognize that input.')
main()

```

Output:
```
Enter 1 for the area
Enter 2 for the circumference
Enter 0 to exit
 0
```